## CSCI 470 Activities and Case Studies

1. For all activities, you are allowed to collaborate with a partner. 
1. For case studies, you should work individually and are **not** allowed to collaborate.

By filling out this notebook and submitting it, you acknowledge that you are aware of the above policies and are agreeing to comply with them.

Some considerations with regard to how these notebooks will be graded:

1. Cells in which "# YOUR CODE HERE" is found are the cells where your graded code should be written.
2. In order to test out or debug your code you may also create notebook cells or edit existing notebook cells other than "# YOUR CODE HERE". We actually highly recommend you do so to gain a better understanding of what is happening. However, during grading, **these changes are ignored**. 
3. You must ensure that all your code for the particular task is available in the cells that say "# YOUR CODE HERE"
4. Every cell that says "# YOUR CODE HERE" is followed by a "raise NotImplementedError". You need to remove that line. During grading, if an error occurs then you will lose points for your work in that section.
5. If your code passes the "assert" statements, then no output will result. If your code fails the "assert" statements, you will get an "AssertionError". Getting an assertion error means you will not receive points for that particular task.
6. If you edit the "assert" statements to make your code pass, they will still fail when they are graded since the autograder will ignore the modified "assert" statement. Make sure you don't edit the assert statements.
7. We may sometimes have "hidden" tests for grading. This means that passing the visible "assert" statements is not sufficient. The "assert" statements are there as a guide but you need to make sure you understand what you're required to do and ensure that you are doing it correctly. Passing the visible tests is necessary but not sufficient to get the grade for that cell.
8. When you are asked to define a function, make sure you **don't** use any variables outside of the parameters passed to the function. You can think of the parameters being passed to the function as a hint. Make sure you're using all of those variables.
9. The **Grading** section at the end of the document (before the **Feedback** section) contains some code for our autograder on GradeScope. You are expected to fail this block of code in your Jupyter environment. DO NOT edit this block of code, or you may not get points for your assignment.
10. Finally, **make sure you run "Kernel > Restart and Run All"** and pass all the asserts before submitting. If you don't restart the kernel, there may be some code that you ran and deleted that is still being used and that was why your asserts were passing.

# Unsupervised Learning - Overview

In this exercise we'll look at data clustering and how we evaluate the results. We'll be using synthetic data with only two features, so we can easily visualize them and see what's happening.

```python
import numpy as np
import pandas as pd
import matplotlib
import matplotlib.pyplot as plt
import sklearn as sk

%matplotlib inline
plt.style.use("ggplot")
```

```python
from sklearn.datasets import make_blobs
from sklearn.model_selection import train_test_split
from sklearn.cluster import KMeans
from sklearn.metrics import adjusted_rand_score, adjusted_mutual_info_score, homogeneity_completeness_v_measure, silhouette_score
```

### Create synthetic data including ground truth class labels

First, we'll create synthetic data with classes that form visually evident clusters, but also some overlap between samples from different ground truth classes.

```python
n = 1000
X, y = make_blobs(n_samples=n, centers=3, n_features=2, random_state=0)
```

```python
plt.scatter(X[:, 0], X[:, 1], marker='o', c=y, s=25, edgecolor='k')
plt.title("3 classes of data")
plt.xlabel("x_1")
plt.ylabel("x_2")
plt.show()
```

```python
X_train, X_test, y_train, y_test = train_test_split(X, y, random_state=0, test_size=0.2)
```

### K-means clustering

Instead of performing classification on this problem, we can withhold the labels from our model and apply a clustering algorithm instead. Let's apply [KMeans](http://scikit-learn.org/stable/modules/generated/sklearn.cluster.KMeans.html) to this example and predict the clusters each point will belong to.

<!-- grade_id: cell-4e365d0ffa423315, points: 0, grade: false, solution: true -->
```python
# Create a model named `model` using KMeans with k = 3
# Fit the model to the training data
# Predict the classes of the test data and save those as `y_pred`
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-bd97498a7f778ba4, points: 5, grade: true, solution: false -->
```python
assert len(model.cluster_centers_) == 3
assert len(y_pred) == len(y_test)
```

We can access the learned cluster centroids using the `cluster_centers_` attribute of the trained model. Let's plot the centroids on top of our data.

```python
centroids = model.cluster_centers_
plt.scatter(X[:, 0], X[:, 1], marker='o', c=y, s=25, edgecolor='k')
plt.scatter(centroids[:, 0], centroids[:, 1], marker="x", c="r")
plt.title("3 classes of data with cluster centroids")
plt.xlabel("X")
plt.ylabel("Y")
plt.show()
```

### Evaluation: Contingency Table

Now that we have our predictions, we can determine how well our model works using a variety of metrics for unsupervised learning. First, let's create the contingency table/matrix for our data. __This is similar to a confusion matrix but the diagonal is irrelevant as the cluster names won't necessarily align with the original classes.__ In addition, the contigency matrix is not necessarily square, because the number of true classes may not be equal to the number of clusters (though they are in this case, since we specified `n_clusters=3`).

The following may be useful:

- Python Sets - [sets](https://docs.python.org/3.6/tutorial/datastructures.html)
- Creating a matrix with all zeros: [np.zeros](https://docs.scipy.org/doc/numpy/reference/generated/numpy.zeros.html)
- Boolean indexing with numpy: [indexing](https://docs.scipy.org/doc/numpy-1.13.0/user/basics.indexing.html)
- Logic operations on numpy arrays: [logic operations](https://docs.scipy.org/doc/numpy-1.13.0/reference/routines.logic.html#logical-operations)

Creating the contingency matrix may not be very straightforward so we've provided a hint for you in case you need it. **We recommend you try to solve it on your own first** but if you're completely stuck, check out the hint below.

Note that the inputs you are given to the function are two sets of sample labelings. Each labeling contains `n` points' labels where each label is the index of the grouping they belong to.

<br>
<details>
<summary>Click here for a hint</summary>
    
There are many approaches to do this and we've outlined 2 approaches that you can think of. The first approach goes through the data and figures out where each sample should be counted in the matrix. The second approach goes through the matrix and determines what the value should be at each location.
    
    
For both approaches, we start by creating an empty matrix:
    
1. Determine the number of unique elements (classes or clusters) of y1 and y2. 
    
2. Initialize a matrix of zeros, with shape equal to (number of unique y1 elements, number of unique y2 elements). 

    
Approach 1: Samples based

1. Iterate over each sample
2. Add 1 to the location in the matrix where that sample belongs
    
Approach 2: Matrix based
    
1. Iterate over each element in the matrix using a pair of nested for loops. 
    
2. For each matrix element, use `np.logical_and()` to help you count the number of y1 elements that fall in that row's class/cluster and the number of y2 elements that fall in that column's class/cluster.
</details>

<!-- grade_id: cell-6c642ff1d5219b74, points: 0, grade: false, solution: true -->
```python
def contingency_table(y_1, y_2):
    """Determine the contingency table/matrix based on two sets of element groupings
    
    Args:
        y_1 (iterable): A labeling of elements into clusters
        y_2 (iterable): Another labeling of elements into clusters
    
    Returns:
        np.ndarray: A matrix with shape n_groups_y1, n_groups_y2 listing the number of 
                    elements from one cluster in y1 that are also in one cluster in y2 
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-5ae5cc99583ff5e1, points: 10, grade: true, solution: false -->
```python
test_cont = contingency_table(np.array([0,0,1,2,2]), np.array([1,1,1,0,2]))
assert test_cont.shape == (3,3)
assert np.all(test_cont.T == contingency_table(np.array([1,1,1,0,2]), np.array([0,0,1,2,2])))
```

```python
contingency_table(y_test, y_pred)
```

Let's examine our model's performance scores based on the other metrics:

```python
print(f"Our model has an adjusted rand index of {adjusted_rand_score(y_test, y_pred):.4f}")
```

```python
print(f"Our model has an adjusted mutual information score of {adjusted_mutual_info_score(y_test, y_pred):.4f}")
```

```python
h,c,v = homogeneity_completeness_v_measure(y_test, y_pred)

print(f"Our model has a homogeneity score of {h:.4f}, completeness score of {c:.4f}, and a v-measure of {v:.4f}.")
```

### Evaluation: Silhoutte Score

Now let's use the silhoutte score and a hyperparameter search on the validation data to determine the best value of k for KMeans. We already expect it to be 3 since we know that is the ground truth however that doesn't mean we are guaranteed the highest score at $k=3$.

The best way to understand the silhoutte score is to examine how it is computed. The treatment at [scikit-learn.org](https://scikit-learn.org/stable/auto_examples/cluster/plot_kmeans_silhouette_analysis.html) is nice, and you can also read the [Wikipedia](https://en.wikipedia.org/wiki/Silhouette_(clustering)) entry.

<!-- grade_id: cell-f95520c4e18c3498, points: 0, grade: false, solution: true -->
```python
def score_ks(ks, features):
    """Get silhoutte scores by applying KMeans to a set of features with a variety of k values 
    
    Args:
        ks (iterable): List of k values to experiment with
        features (iterable): The features to train the model on
    
    Returns:
        iterable: A list of scores of models trained using each value of k
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
ks = [2,3,4,5,6,10]
scores = score_ks(ks, X)
```

<!-- grade_id: cell-2cb9ba46d378fd5a, points: 5, grade: true, solution: false -->
```python
assert len(scores) == len(ks)
assert scores[1] == max(scores)
```

```python
maxScore = 0
bestK = ks[0]
for k, score in zip(ks, scores):
    if score > maxScore:
        bestK = k
        maxScore = score
    print(f"With k = {k}, the silhoute score is {score:.4f}")
    
print(f"\n\nThe best k value is k={bestK}, scoring {maxScore:.4f}.")
```

```python
plt.plot(ks, scores)
plt.xlabel("k")
plt.ylabel("silhoutte score")
plt.title("Silhoutte scores of KMeans clustering")
plt.show()
```

### What is the "right" value for k in a given K-means clustering problem?

When we have unlabeled data, there may be no "right" value for k. We thus do our best, using a heuristic measure that helps us determine a "good" value. That is the purpose served by the silhouette score. Try changing the number of blobs generated in the data and see if the silhouette score always give us a value for k that is equal to the number of ground truth classes in our synthetic data. It won't!

__Also observe how the contingency matrix looks when you try different values of clusters versus classes.__

# Grading
The following code block is purely used for grading. If you find any error, you can ignore. DO NOT MODIFY THE CODE BLOCK BELOW.

```python
# Autograding with Otter Grader
import otter
grader = otter.Notebook()
grader.check_all()
```

## Feedback

<!-- grade_id: feedback, points: 0, grade: false, solution: true -->
```python
def feedback():
    """Provide feedback on the contents of this exercise
    
    Returns:
        string
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: feedback-tests, points: 0, grade: true, solution: false -->
```python

```